Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

```python
NAME = 'Paul Natland'
UMICH_UNIQNAME = 'pnatland'
```

```python
"""
ACKNOWLEDGEMENTS: For this assignment, I benefited from lecture videos, searches on Stacks Overflow, and conversations on SLACK
"""
```

Output:
```
'\nACKNOWLEDGEMENTS: For this assignment, I benefited from lecture videos, searches on Stacks Overflow, and conversations on SLACK\n'
```

---

```python
version = "v2.3.060120"
```

# SIADS 515 Week 2 Homework (HW2)

For this assignment you'll be using a Python class we've written that provides data from the `metrics.csv` and `position.csv` files in the `assets` directory. You will not need to directly access those files to complete the assingment, but you may want to see their contents.

```python
from assets.api import BikeRideApi
api = BikeRideApi('assets/metrics.csv', 'assets/position.csv')
```

This class provides a `get_data()` method.

You can use the built-in `help()` function to read the docstring for `get_data()` ...

```python
help(api.get_data)
```

Output:
```
Help on method get_data in module assets.api:

get_data(dataset_name, offset, count=5) method of assets.api.BikeRideApi instance
    Returns a tuple of "count" tuples starting from "offset" for the
    given "dataset"
    
    Arguments:
        dataset_name {str} -- The dataset to pull data from ("metrics"
            or "position")
        offset {int} -- Offset from the beginning of the dataset (must
            be greater than 0)
    
    Keyword Arguments:
        count {int} -- Number of rows to include (default: 5; must be 
            between 1 and 10 inclusive)
    
    Raises:
        ValueError: Invalid parameter values
    
    Returns:
        tuple -- A tuple of "count" tuples starting from "offset" for the
            given "dataset". Returns an empty tuple when there is no more 
            data.
```

__Exploration__

```python
api.get_data('metrics', offset=0)
```

Output:
```
((10.87, 285.79999999999995, 1.773),
 (71.85, 285.0, 5.533),
 (108.02, 284.0, 6.485),
 (170.23, 284.0, 6.951),
 (229.27, 285.0, 6.224))
```

```python
api.get_data('position', offset=0)
```

Output:
```
((504719750, -998493490),
 (504717676, -998501870),
 (504716354, -998506792),
 (504714055, -998515244),
 (504711900, -998523278))
```

```python
api.get_data('position', offset=0)[2][0]
```

Output:
```
504716354
```

```python
api.get_data('metrics', offset=4, count=1)
```

Output:
```
((229.27, 285.0, 6.224),)
```

```python
#not desired format
#list(zip(api.get_data('metrics', offset=0), api.get_data('position', offset=0)))
```

```python
n = 0
while True:
    try:
        api.get_data('position', offset=n)[0]
        n += 1
    except IndexError:
        break
print(n)
```

Output:
```
207
```

```python
#this is the desired format in tuple form
api.get_data('metrics', offset=0)[0] + api.get_data('position', offset=0)[0]
```

Output:
```
(10.87, 285.79999999999995, 1.773, 504719750, -998493490)
```

---
**NOTE** 

The `get_data()` method artificially adds a delay of 1 millisecond to each call. This mimics the latency that is common when fetching data from network sources like an HTTP API. In the real world, 1 millisecond would be very low latency for an HTTP API, but it adds enough of a delay that you'll be able to see very noticable differences depending on how you optimize your code for efficency.

---
## Question 1
Write a **generator** called `combine_data()` that meets these requirements:
  - It doesn't take any parameters.
  - It must use the `get_data()` method descibed above.
  - It must use a value for the "count" parameter greater than 1.
  - It must use only the Python standard library (e.g. no IPython magic commands and no external libraries like Pandas)
  - It must yield a combined structure for each row in the "metrics" and "position" datasets.
  - Each yielded item must be an iterable (e.g. a list or a tuple is fine), that contains, in this order, the values for:
    - "Distance"
    - "Altitude"
    - "Speed"
    - "Latitude"
    - "Longitude"
  - It must stop yielding when either dataset runs out of data.
  - It must handle cases other than the data in the `assets` directory (e.g. don't hard code dataset lengths, or other values).
  
Note:
  - The `zip` built-in function could be useful in your solution, but is not required.

<!-- grade_id: cell-a187bd9c83d3f08b, points: 0, grade: false, solution: true -->
```python
def combine_data():
    num = 0
    while True:
        try:
            result = api.get_data('metrics', offset=num, count = 2)[0] + api.get_data('position', offset=num, count = 2)[0]
            num += 1
        except IndexError:
            break
        yield result
```

```python
%%timeit
gen = combine_data()
```

Output:
```
210 ns ± 2.01 ns per loop (mean ± std. dev. of 7 runs, 1000000 loops each)
```

```python
%%timeit
gen = combine_data()

list(gen)
```

Output:
```
477 ms ± 1.24 ms per loop (mean ± std. dev. of 7 runs, 1 loop each)
```

<!-- grade_id: cell-7b8a95d40b760538, points: 1, grade: true, solution: false -->
```python
# Question 1, Check 1
# 1 point

gen = combine_data()

# Make sure your solution is a generator
from types import GeneratorType
assert type(gen) == GeneratorType, 'wrong type, should be a generator'

# There are no hidden tests in this cell.
```

<!-- grade_id: cell-bc023ea9d5ff26ca, points: 1, grade: true, solution: false -->
```python
# Question 1, Check 2 
# 1 point

gen = combine_data()

# Make sure you stop yielding when one dataset ends 
assert len(list(gen)) == 207, 'wrong number of results'

# There are no hidden tests in this cell.
```

<!-- grade_id: cell-ab0589a1a1ff0854, points: 1, grade: true, solution: false -->
```python
# Question 1, Check 3
# 1 point.

gen = combine_data()

# Make sure your solution yields the correct 1st row
assert next(gen) == (10.87, 285.79999999999995, 1.773, 504719750, -998493490), 'incorrect data'

# Make sure your solution yields the correct 2nd row
assert next(gen) == (71.85, 285.0, 5.533, 504717676, -998501870), 'incorrect data'

# Make sure your solution yields the correct 3rd row
assert next(gen) == (108.02, 284.0, 6.485, 504716354, -998506792), 'incorrect data'

# There are no hidden tests in this cell.
```

<!-- grade_id: cell-5891a33ffe23f717, points: 1, grade: true, solution: false -->
```python
# Question 1, Check 4
# 1 point

# There are hidden tests in this cell to check that get_data() is called
# NOTE: To get this check to pass, your solution just needs to call get_data()
```

<!-- grade_id: cell-ba7c1a8f9e8160c2, points: 1, grade: true, solution: false -->
```python
# Question 1, Check 5
# 1 point.

# NOTE: To get this test to pass, your solution needs to either not specify the "count"
# param, or specifiy a value of greater than 1 for the "count" param.

# There are hidden tests in this cell to check that get_data() is called with "count" > 1
```

<!-- grade_id: cell-07c5c64ac223bf52, points: 5, grade: true, solution: false -->
```python
# Question 1, Check 6
# 5 points

# There are hidden tests in this cell to check that at least 90% rows are correct for the given data.
```

<!-- grade_id: cell-9dbe37c3bb9a9f79, points: 5, grade: true, solution: false -->
```python
# Question 1, Check 7
# 5 points

# There are hidden tests in this cell to check that at 100% rows are correct for the given data.
```

<!-- grade_id: cell-0dae3f7986460f1a, points: 5, grade: true, solution: false -->
```python
# Question 1, Check 8
# 5 points

# There are hidden tests in this cell to check the solution with other input cases.
```

---
## Question 2
Write a **function** called `get_speed()` that meets these requirements:
  - It must take 1 parameter: an integer that represents the row of data from which to return the speed.
  - It must use the `get_data()` method descibed above.
  - It must use only the Python standard library (e.g. no IPython magic commands and no external libraries like Pandas)
  - It must be decorated with the `lru_cache` decorator from the `functool` package in the standard library.
  - It must return an numeric value for the speed at the given data index.
  - It must return `None` if the data at the given index is empty (past the end of the dataset)
  - It must handle cases other than the data in the `assets` directory (e.g. don't hard code dataset lengths, or other values).
  
Note:
  - It is probably simplest to use a value for the "count" parameter of 1 (unlike the requirement in Question 1).

```python
#gets the speed for row = offset
api.get_data('metrics', offset=4, count=1)[0][2]
```

Output:
```
6.224
```

### solution

<!-- grade_id: cell-86088d3099578a9f, points: 0, grade: false, solution: true -->
```python
from functools import lru_cache

@lru_cache(16)
def get_speed(i):
    # i equals the row number
    while True:
        try:
            result = api.get_data('metrics', offset=i, count=1)[0][2]
        except IndexError:
            return None
        return result
```

The `%timeit` magic command is a quick way to judge the speed of a function.

```python
%timeit get_speed(0)
```

Output:
```
74.3 ns ± 0.912 ns per loop (mean ± std. dev. of 7 runs, 10000000 loops each)
```

The `lru_cache` decorator actually turns `get_speed` into a callable object so it still works like the original function, but also has some helpful properities. For example the `__wrapped__` property is the oringal function, so we can see how much worse its performance is by calling that with `%timeit` instead.

This call should be significantly slower than the one above. Given that the underlying method is adding a 1 millisecond delay, the uncached version of `get_speed()` should take at least 1 millisecond.

```python
# This cell will fail if get_speed() is not decorated with lru_cache

%timeit get_speed.__wrapped__(0)
```

Output:
```
1.14 ms ± 3.29 µs per loop (mean ± std. dev. of 7 runs, 1000 loops each)
```

<!-- grade_id: cell-72816bab60b45bd8, points: 1, grade: true, solution: false -->
```python
# Question 2, Check 1
# 1 point.

# Make sure your solution uses the lru_cache decorator
from functools import _lru_cache_wrapper
assert type(get_speed) == _lru_cache_wrapper, 'solution must be decorated with the lru_cache decorator'

# There are no hidden tests in this cell.
```

<!-- grade_id: cell-c115d45e02393b0b, points: 1, grade: true, solution: false -->
```python
# Question 2, Check 2
# 1 point

# Make sure your solution returns the correct values for the first 3 indexes
assert get_speed(0) == 1.773, 'wrong data value'
assert get_speed(1) == 5.533, 'wrong data value'
assert get_speed(2) == 6.485, 'wrong data value'

# There are no hidden tests in this cell.
```

<!-- grade_id: cell-3c0b8f6483a84020, points: 1, grade: true, solution: false -->
```python
# Question 2, Check 3
# 1 point

# Make sure your solution returns the None for indexes past the range of the dataset
assert get_speed(214) == None, 'must return None for indexes past the range of the dataset'
assert get_speed(1000) == None, 'must return None for indexes past the range of the dataset'

# There are no hidden tests in this cell.
```

<!-- grade_id: cell-80ba5d93436d25ee, points: 4, grade: true, solution: false -->
```python
# Question 2, Check 4
# 4 points

# There are hidden tests in this cell to check the solution returns at least 90% of the correct results.
```

<!-- grade_id: cell-b7389411925d4b8a, points: 4, grade: true, solution: false -->
```python
# Question 2, Check 5
# 4 points

# There are hidden tests in this cell to check the solution returns 100% of the correct results.
```

<!-- grade_id: cell-05d07679deb8c897, points: 4, grade: true, solution: false -->
```python
# Question 2, Check 6
# 4 points

# There are hidden tests in this cell to check the solution with other input cases.
```

---
## Question 2b
Write a **generator** called `speed_rolling_average()` that meets these requirements:
  - It must take 1 parameter: an integer that represents the number of rows over which to calculate a rolling average. This must default to 5.
  - It must use the `get_speed()` method you implemented above for Question 2.
  - It must use only the Python standard library (e.g. no IPython magic commands and no external libraries like Pandas)
  - It must yield an numeric value for the average speed over the previous number of rows as defined by the `period` parameter.
  - It must yield `None` if there have been fewer than `period` number of rows so far in the dataset.
  - It must handle cases other than the data in the `assets` directory (e.g. don't hard code dataset lengths, or other values).

### my better solution

```python
def speed_rolling_average(period=5):
    # the parameter value represents the number of rows to consider for the rolling average
    n = 0
    while n < (period-1):
        yield None
        n+= 1
    i = 0
    while True:
        try:
            t = 0
            speeds_to_average = []
            while t < period:
                value = get_speed(i+t)
                speeds_to_average.append(value)
                t += 1
            yield sum(speeds_to_average)/len(speeds_to_average)
            i += 1
        except TypeError:
            break
```

```python
%%timeit
gen = speed_rolling_average(5)

list(gen)
```

Output:
```
247 ms ± 929 µs per loop (mean ± std. dev. of 7 runs, 1 loop each)
```

### my first solution

```python
def speed_rolling_average(period=5):
    # the parameter value represents the number of rows to consider for the rolling average
    n = 0
    while True:
        try:
            api.get_data('metrics', offset=n)[0]
            n += 1
        except IndexError:
            break
    for i in range(n):
        if i < (period-1):
            yield None
        else:
            break
    for i in range(n):
        try:
            t = 0
            speeds_to_average = []
            while t < period:
                value = get_speed(i+t)
                speeds_to_average.append(value)
                t += 1
            yield sum(speeds_to_average)/len(speeds_to_average)
        except TypeError:
            break
```

```python
%%timeit
gen = speed_rolling_average(5)

list(gen)
```

Output:
```
494 ms ± 2.61 ms per loop (mean ± std. dev. of 7 runs, 1 loop each)
```

<!-- grade_id: cell-6491afffc16e3cc2, points: 0, grade: false, solution: true -->
```python
gen = speed_rolling_average(5)

len(list(gen))
```

Output:
```
214
```

<!-- grade_id: cell-6e87683e2f331c7a, points: 1, grade: true, solution: false -->
```python
# Question 2b, Check 1
# 1 point.

gen = speed_rolling_average(5)

# Make sure your solution is a generator
from types import GeneratorType
assert type(gen) == GeneratorType, 'wrong type, should be a generator'

# There are no hidden tests in this cell.
```

<!-- grade_id: cell-9f86de205d0cec1e, points: 1, grade: true, solution: false -->
```python
# Question 2b, Check 2 
# 1 point

gen = speed_rolling_average(5)

# Make sure you stop yielding when one dataset ends 
assert len(list(gen)) == 214, 'wrong number of results'

# There are no hidden tests in this cell.
```

<!-- grade_id: cell-44de658595075419, points: 1, grade: true, solution: false -->
```python
# Question 2b, Check 3
# 1 point.

gen = speed_rolling_average(5)

# Make sure your solution yields None for period-1 iterations
assert next(gen) == None, 'should yield None for i < period'
assert next(gen) == None, 'should yield None for i < period'
assert next(gen) == None, 'should yield None for i < period'
assert next(gen) == None, 'should yield None for i < period'

# There are no hidden autograder tests in this cell.
```

<!-- grade_id: cell-6e48c7fb42457bf1, points: 2, grade: true, solution: false -->
```python
# Question 2b, Check 4
# 1 point.

# Make sure your solution yields the correct first rolling average speed
assert next(gen) == 5.3932, 'wrong data result'

# Make sure your solution yields the correct second rolling average speed
assert next(gen) == 6.1508, 'wrong data result'

# There are no hidden tests in this cell.
```

<!-- grade_id: cell-0a497e5b983da4d0, points: 5, grade: true, solution: false -->
```python
# Question 2b, Check 5
# 5 points

# There are hidden tests in this cell to check the solution returns at least 90% of the correct results.
```

<!-- grade_id: cell-8eaaf3f7099d8fa5, points: 5, grade: true, solution: false -->
```python
# Question 2b, Check 6
# 5 points

# There are hidden tests in this cell to check the solution returns 100% of the correct results.
```

```python

```