Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

```python
# Please add your name and University of Michagan uniqname here...

NAME = 'Paul Natland'
UMICH_UNIQNAME = 'pnatland'
```

```python
"""
ACKNOWLEDGEMENTS: For this assignment, I benefited from lecture videos, searches on Stacks Overflow, and conversations on SLACK
"""
```

Output:
```
'\nACKNOWLEDGEMENTS: For this assignment, I benefited from lecture videos, searches on Stacks Overflow, and conversations on SLACK\n'
```

---

```python
version = "v2.3.060120"
```

# SIADS 515 Week 4 Homework (HW4)

## Background

The motivation for this assignment is to **improve the efficiency of code** that finds the similarity between any two given documents.  There are many ways to calculate similarity (or distance) between two documents, but the most effective way is to represent each document as a multi-dimensional vector where each dimension corresponds to a word, and the value along that dimension is the number of times that word occurs in a document.  Let's take a look at a simplified case where we only have two dimensions:

![](assets/CosineDistanceSimilarity.png)

In the above diagram, Item 1 and Item 2 refer to two documents.  The angle between them (ð›³) can range from -180 to +180 degrees.  The cosine of angles, in this case, has a nice property in that the cosine of an angle of 0 degrees is 1, the cosine of an angle of 90 degrees is 0, and the cosine of an angle of 180 degrees is -1.  In other words, the cosine of the angle between the vector representation of a document behaves much like a correlation coefficient.  A cosine of 1 indicates the documents are identical, a cosine of 0 indicates the documents are independent of each other, and a cosine of -1 indicates the documents are "opposites" (note: the latter requires a specific setup that is beyond the scope of this course).

Your task, for this assignment, is to improve the efficiency of the code.

## Steps

1. Run the notebook as it is.
1. Study each function in the code, making notes to yourself in preparation for the next step.  
1. Annotate each function with a detailed description of what the input and output parameters are.  Provide a brief description, in your own words, of what each function does.  Use docstrings ([PEP-257](https://www.python.org/dev/peps/pep-0257/)) to document the code.
1. Make changes to the code to improve its efficiency. 

Note
  - You may [refactor](https://en.wikipedia.org/wiki/Code_refactoring) to combine functions or make other changes if that helps.

```python
# docdist1.py
# Original author: Ronald L. Rivest
# Some adaptation by Chris Teplovs and Kevyn Collins-Thompson
#
#
# This program computes the "distance" between two text files
# as the angle between their word frequency vectors.
#
# For each input file, a word-frequency vector is computed as follows:
#    (1) the specified file is read in
#    (2) it is converted into a list of alphanumeric "words"
#        Here a "word" is a sequence of consecutive alphanumeric
#        characters.  Non-alphanumeric characters are treated as blanks.
#        Case is not significant.
#    (3) for each word, its frequency of occurrence is determined
#    (4) the word/frequency lists are sorted into order alphabetically
#
# The "distance" between two vectors is the angle between them.
# If x = (x1, x2, ..., xn) is the first vector (xi = freq of word i)
# and y = (y1, y2, ..., yn) is the second vector,
# then the angle between them is defined as:
#    d(x,y) = arccos(inner_product(x,y) / (norm(x)*norm(y)))
# where:
#    inner_product(x,y) = x1*y1 + x2*y2 + ... xn*yn
#    norm(x) = sqrt(inner_product(x,x))
```

```python
import math
    # math.acos(x) is the arccosine of x.
    # math.sqrt(x) is the square root of x.

import string
    # string.join(words,sep) takes a given list of words,
    #    and returns a single string resulting from concatenating them
    #    together, separated by the string sep .
    # string.lower(word) converts word to lower-case

import sys
    # sys.exit() allows us to quit (if we can't read a file)

"""New Imports"""
from functools import lru_cache
    #to use the @lru_cache decorator...all over the place!
    
import re
    #to apply regular expressions to extract words from text (in my read_file_new function)

from collections import defaultdict
    #for better dictionary functionality (in my read_file_new function)
```

```python
%load_ext line_profiler
```

## These are the code originals with updated docstrings according to PEP-257

```python
# this is the original code with updated docstring

def read_file(filename):
    """Read the text file (filename) and return a list of the lines from it."""
    try:
        fp = open(filename)
        L = fp.readlines()
    except IOError as excObj:
        print(str(excObj))
        print("Error opening or reading input file: " + filename)
        sys.exit()
    return L
```

```python
repr(read_file.__doc__)
```

Output:
```
"'Read the text file (filename) and return a list of the lines from it.'"
```

```python
# this is the original code with updated docstring

def get_words_from_line_list(L):  
    """Parse the given list (L) of text lines and return a list of all words found."""
    word_list = []
    for line in L:
        words_in_line = get_words_from_string(line)
        word_list = word_list + words_in_line
    return word_list
```

```python
repr(get_words_from_line_list.__doc__)
```

Output:
```
"'Parse the given list (L) of text lines and return a list of all words found.'"
```

```python
#this is the original code with updated docstring

def get_words_from_string(line):
    """Parse the give the string (line) and return a list of words in it after converting all words to lower-case."""
    word_list = []          # accumulates words in line
    character_list = []     # accumulates characters in word
    for c in line:
        if c.isalnum():
            character_list.append(c)
        elif len(character_list)>0:
            word = str.join("", character_list)
            word = str.lower(word)
            word_list.append(word)
            character_list = []
    if len(character_list)>0:
        word = str.join("", character_list)
        word = str.lower(word)
        word_list.append(word)
    return word_list
```

```python
repr(get_words_from_string.__doc__)
```

Output:
```
"'Parse the give the string (line) and return a list of words in it after converting all words to lower-case.'"
```

```python
#adjustment based on Part 2 of assignment, using .extend()
# def get_words_from_line_list2(L):  
#     """using extend """
#     ### Parse the given list L of text lines into words.
#     ### Return list of all words found.

#     word_list = []
#     for line in L:
#         words_in_line = get_words_from_string(line)
#         word_list.extend(words_in_line)
#     return word_list
```

```python
#this is the original code with updated docstring

def count_frequency(word_list):
    """Parse the given list of words in word_list and return a list of lists containing the word and its frequency."""
    L = []
    for new_word in word_list:
        for entry in L:
            if new_word == entry[0]:
                entry[1] = entry[1] + 1
                break
        else:
            L.append([new_word,1])
    return L
```

```python
repr(count_frequency.__doc__)
```

Output:
```
"'Parse the given list of words in word_list and return a list of lists containing the word and its frequency.'"
```

```python
# code update based on part 2 of assignment
# def count_frequency2(word_list):
        
#     ### Return a list giving pairs of form: (word,frequency)

#     D = {}
#     for new_word in word_list:
#         for entry in D:
#             if new_word in D:
#                 D[new_word] += 1
#                 break
#         else:
#             D[new_word] = 1
#     return D
```

```python
#this is the original code with updated docstring

def insertion_sort(A):
    """Return the given list (A) sorted in place according to the item of index == 0.

    From Cormen/Leiserson/Rivest/Stein, Introduction to Algorithms (second edition), page 17, 
    modified to adjust for fact that Python arrays use 0-indexing.
    """
    for j in range(len(A)):
        key = A[j]
        # insert A[j] into sorted sequence A[0..j-1]
        i = j-1
        while i>-1 and A[i]>key:
            A[i+1] = A[i]
            i = i-1
        A[i+1] = key
    return A
```

```python
repr(insertion_sort.__doc__)
```

Output:
```
"'Return the given list (A) sorted in place according to the item of index == 0.\\n\\n    From Cormen/Leiserson/Rivest/Stein, Introduction to Algorithms (second edition), page 17, \\n    modified to adjust for fact that Python arrays use 0-indexing.\\n    '"
```

```python
#this is the original code with updated docstring

def word_frequencies_for_file(filename,verbose=False):
    """For the give text file (filename), return an alphabetically sorted list of words contained and their frequency.
    
    Returns An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]
    The function also prints a statement including the number of lines, the number of words, and the number of 
    distinct words in the text file (filename).
    """
    line_list = read_file(filename)
    word_list = get_words_from_line_list(line_list)
    freq_mapping = count_frequency(word_list)
    insertion_sort(freq_mapping)
    if verbose:
        print("File",filename,":", len(line_list),"lines,", len(word_list),"words,", len(freq_mapping),"distinct words")

    return freq_mapping
```

```python
repr(word_frequencies_for_file.__doc__)
```

Output:
```
"'For the give text file (filename), return an alphabetically sorted list of words contained and their frequency.\\n    \\n    Returns An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]\\n    The function also prints a statement including the number of lines, the number of words, and the number of \\n    distinct words in the text file (filename).\\n    '"
```

```python
#this is the original code --- docstring needs updating

def inner_product(L1,L2):
    """Return an inner product of the two vectors provided (L1, L2).
    
    Args:
        L1: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]
        L2: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]
    """
    ### Inner product between two vectors, where vectors
    ### are represented as alphabetically sorted (word,freq) pairs.

    ### Example: inner_product([["and",3],["of",2],["the",5]],
    ###                        [["and",4],["in",1],["of",1],["this",2]]) = 14.0

    sum = 0.0
    i = 0
    j = 0
    while i<len(L1) and j<len(L2):
        # L1[i:] and L2[j:] yet to be processed
        if L1[i][0] == L2[j][0]:
            # both vectors have this word
            sum += L1[i][1] * L2[j][1]
            i += 1
            j += 1
        elif L1[i][0] < L2[j][0]:
            # word L1[i][0] is in L1 but not L2
            i += 1
        else:
            # word L2[j][0] is in L2 but not L1
            j += 1
    return sum
```

```python
repr(inner_product.__doc__)
```

Output:
```
"'Return an inner product of the two vectors provided (L1, L2).\\n    \\n    Args:\\n        L1: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]\\n        L2: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]\\n    '"
```

```python
#this is the original code  --- docstring needs updating

def vector_angle(L1,L2):
    """Returns the angle (in radians) between the given vectors (L1 and L2).
    
    The function performs a dot product of the two vectors and returns the angle between them.
    
    Args:
        L1: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]
        L2: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]
    """
    numerator = inner_product(L1,L2)
    denominator = math.sqrt(inner_product(L1,L1)*inner_product(L2,L2))
    return numerator/denominator
```

```python
repr(vector_angle.__doc__)
```

Output:
```
"'Returns the angle (in radians) between the given vectors (L1 and L2).\\n    \\n    The function performs a dot product of the two vectors and returns the angle between them.\\n    \\n    Args:\\n        L1: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]\\n        L2: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]\\n    '"
```

```python
#this is the original code  ---- just do updated docstring
def document_similarity(filename_1, filename_2, verbose=True):
    """For each of the input text files, find the distinct words present and their frequency, vectorize that data
    and return the angle and and cosine of the vectors.
    
    The function takes as input two text files, counts the word frequency in each one, vectorizes that information
    and calculates the Cosine Similarity of the two vectors produced.  For each input text file (filename_1, filename_2)
    the function prints the number of lines, words, and distinct words in each and, if verbose=True, also
    prints the cosine and angle for the vectors described.
    """
    sorted_word_list_1 = word_frequencies_for_file(filename_1, verbose)
    sorted_word_list_2 = word_frequencies_for_file(filename_2, verbose)
    cosine = vector_angle(sorted_word_list_1,sorted_word_list_2)
    # Use f-strings; see https://realpython.com/python-f-strings/ for more information
    if verbose:
        print(f"The cosine between the documents is {cosine : 0.6f}.")
        print(f"The angle between the documents is {math.acos(cosine) : 0.6f} radians or {math.acos(cosine)*180/math.pi : .0f} degrees.")
```

```python
repr(document_similarity.__doc__)
```

Output:
```
"'For each of the input text files, find the distinct words present and their frequency, vectorize that data\\n    and return the angle and and cosine of the vectors.\\n    \\n    The function takes as input two text files, counts the word frequency in each one, vectorizes that information\\n    and calculates the Cosine Similarity of the two vectors produced.  For each input text file (filename_1, filename_2)\\n    the function prints the number of lines, words, and distinct words in each and, if verbose=True, also\\n    prints the cosine and angle for the vectors described.\\n    '"
```

```python
document_similarity('assets/short.t1.txt','assets/short.t2.txt')
```

Output:
```
File assets/short.t1.txt : 200 lines, 1855 words, 772 distinct words
File assets/short.t2.txt : 200 lines, 752 words, 334 distinct words
The cosine between the documents is  0.712533.
The angle between the documents is  0.777694 radians or  45 degrees.
```

## Here is my refactored code to yield the same product

```python
@lru_cache(128)
def read_file_new(filename):
    """Read the text file (filename) and return a dictionary (defaultdict) with unique words and their frequency.
    
    Function uses regex to extract all of the words in each line of filename after converting words to lower-case.
    It then stores the words as keys and counts their frequency in the defaultdict.
    """
    try:
        wordfreq = defaultdict(int)
        with open(filename, 'r') as file:
            for line in file:
                for word in re.findall(r'[A-Za-z0-9]+', line.lower()):
                    wordfreq[word]+=1
        
    except IOError as excObj:
        print(str(excObj))
        print("Error opening or reading input file: " + filename)
        sys.exit()
        
    return wordfreq
```

```python
repr(read_file_new.__doc__)
```

Output:
```
"'Read the text file (filename) and return a dictionary (defaultdict) with unique words and their frequency.\\n    \\n    Function uses regex to extract all of the words in each line of filename after converting words to lower-case.\\n    It then stores the words as keys and counts their frequency in the defaultdict.\\n    '"
```

```python
%timeit read_file_new('assets/short.t1.txt')
```

Output:
```
72.7 ns Â± 0.614 ns per loop (mean Â± std. dev. of 7 runs, 10000000 loops each)
```

```python
%timeit read_file_new.__wrapped__('assets/short.t1.txt')
```

Output:
```
7.01 ms Â± 99.7 Âµs per loop (mean Â± std. dev. of 7 runs, 100 loops each)
```

```python
#this is the original code, just used the decorator @lru_cache

@lru_cache(128)
def inner_product(L1,L2):
    """Return an inner product of the two vectors provided (L1, L2).
    
    Args:
        L1: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]
        L2: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]
    """
    sum = 0.0
    i = 0
    j = 0
    while i<len(L1) and j<len(L2):
        # L1[i:] and L2[j:] yet to be processed
        if L1[i][0] == L2[j][0]:
            # both vectors have this word
            sum += L1[i][1] * L2[j][1]
            i += 1
            j += 1
        elif L1[i][0] < L2[j][0]:
            # word L1[i][0] is in L1 but not L2
            i += 1
        else:
            # word L2[j][0] is in L2 but not L1
            j += 1
    return sum
```

```python
#this is the original code, just used the decorator @lru_cache

@lru_cache(128)
def vector_angle(L1,L2):
    """Returns the angle (in radians) between the given vectors (L1 and L2).
    
    The function performs a dot product of the two vectors and returns the angle between them.
    
    Args:
        L1: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]
        L2: An alphabetically sorted list of lists that represent word frequency pairs in the format [word, freq]
    """
    numerator = inner_product(L1,L2)
    denominator = math.sqrt(inner_product(L1,L1)*inner_product(L2,L2))
    return numerator/denominator
```

```python
@lru_cache(128)
def document_similarity_new(filename_1, filename_2, verbose=True):
    """For each of the input text files, find the distinct words present and their frequency, vectorize that data
    and return the angle and and cosine of the vectors.
    
    The function takes as input two text files, counts the word frequency in each one, vectorizes that information
    and calculates the Cosine Similarity of the two vectors produced.  For each input text file (filename_1, filename_2)
    the function prints the number of lines, words, and distinct words in each and, if verbose=True, also
    prints the cosine and angle for the vectors described.
    """
    list1 = [(k,v) for k,v in read_file_new(filename_1).items()]
    list2 = [(k,v) for k,v in read_file_new(filename_2).items()]
    list1.sort()
    list2.sort()
    
    list1_tuple = tuple(list1)
    list2_tuple = tuple(list2)
    
    cosine = vector_angle(list1_tuple,list2_tuple)

    if verbose:
        print(f"The cosine between the documents is {cosine : 0.6f}.")
        print(f"The angle between the documents is {math.acos(cosine) : 0.6f} radians or{math.acos(cosine)*180/math.pi : .0f} degrees.")
```

## Returns the same cosine and angle as the original function

```python
document_similarity_new('assets/short.t1.txt','assets/short.t2.txt', verbose=True)
```

Output:
```
The cosine between the documents is  0.712533.
The angle between the documents is  0.777694 radians or 45 degrees.
```

---

### Perfomance of short text cases

```python
%timeit document_similarity_new('assets/short.t1.txt','assets/short.t2.txt', verbose=False)
```

Output:
```
201 ns Â± 1.92 ns per loop (mean Â± std. dev. of 7 runs, 1000000 loops each)
```

```python
%timeit document_similarity_new.__wrapped__('assets/short.t1.txt','assets/short.t2.txt', verbose=False)
```

Output:
```
388 Âµs Â± 5.91 Âµs per loop (mean Â± std. dev. of 7 runs, 1000 loops each)
```

### Performance of long text cases

```python
%timeit document_similarity_new('assets/t1.verne.txt', 'assets/t2.bobsey.txt', verbose=False)
```

Output:
```
201 ns Â± 1.11 ns per loop (mean Â± std. dev. of 7 runs, 10000000 loops each)
```

### Perfomance of very long text cases

```python
%timeit document_similarity_new('assets/t5.churchill.txt', 'assets/t8.shakespeare.txt', verbose=False)
```

Output:
```
235 ns Â± 3.2 ns per loop (mean Â± std. dev. of 7 runs, 1000000 loops each)
```

```python
%timeit document_similarity_new.__wrapped__('assets/t5.churchill.txt', 'assets/t8.shakespeare.txt', verbose=False)
```

Output:
```
41.9 ms Â± 556 Âµs per loop (mean Â± std. dev. of 7 runs, 10 loops each)
```

## I adjusted the assert statments because of my refactored code

```python
# Result validation
# 20 points

# This check verifies that your results are correct for the short text cases.
list1 = [(k,v) for k,v in read_file_new('assets/short.t1.txt').items()]
list2 = [(k,v) for k,v in read_file_new('assets/short.t2.txt').items()]
list1.sort()
list2.sort()

list1_tuple = tuple(list1)
list2_tuple = tuple(list2)

sorted_word_list_1 = list1_tuple
sorted_word_list_2 = list2_tuple
cosine = vector_angle(sorted_word_list_1,sorted_word_list_2)

def close_match(a, b):
    return round(float(a), 3) == round(float(b), 3)

assert len(sorted_word_list_1) == 772
assert len(sorted_word_list_2) == 334

assert close_match(vector_angle(sorted_word_list_1,sorted_word_list_2), 0.713)
assert close_match(math.acos(cosine), 0.778)  # Documents angle in radians
assert close_match(math.acos(cosine)*180/math.pi, 44.559)  # Documents angle in degrees

# There are no hidden autograder tests in this cell.
```

## In office hours and on SLACK I heard and read (I believe) that it is okay to rewrite or even ignore what is below if it is no longer relevant.

<!-- grade_id: cell-1a-63aea205c14ed, points: 20, grade: true, solution: false -->
```python
# Result validation
# 20 points

# This check verifies that your results are correct for the short text cases.

sorted_word_list_1 = word_frequencies_for_file('assets/short.t1.txt', verbose=False)
sorted_word_list_2 = word_frequencies_for_file('assets/short.t2.txt', verbose=False)
cosine = vector_angle(sorted_word_list_1,sorted_word_list_2)

def close_match(a, b):
    return round(float(a), 3) == round(float(b), 3)

assert len(sorted_word_list_1) == 772
assert len(sorted_word_list_2) == 334

assert close_match(vector_angle(sorted_word_list_1,sorted_word_list_2), 0.713)
assert close_match(math.acos(cosine), 0.778)  # Documents angle in radians
assert close_match(math.acos(cosine)*180/math.pi, 44.559)  # Documents angle in degrees

# There are no hidden autograder tests in this cell.
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-42-692d048cfedf> in <module>
      6 sorted_word_list_1 = word_frequencies_for_file('assets/short.t1.txt', verbose=False)
      7 sorted_word_list_2 = word_frequencies_for_file('assets/short.t2.txt', verbose=False)
----> 8 cosine = vector_angle(sorted_word_list_1,sorted_word_list_2)
      9 
     10 def close_match(a, b):

TypeError: unhashable type: 'list'
```